## Задача 1 - простой температурный калькулятор

В первой задаче ваша цель состоит в том, чтобы создать функцию, которая преобразует входную температуру из градусов по Фаренгейту в градусы по Цельсию. Формула преобразования из Фаренгейта в Цельсия можно найти ниже.

\begin{equation}
  \Large
  T_{\mathrm{Celsius}} = \left( T_{\mathrm{Fahrenheit}} - 32 \right) / 1.8
\end{equation}

Если все в вашем скрипте работает должным образом, должен работать следующий тестовый пример:

```python
>>> print(f"32 градуса по Фаренгейту равны {fahr_to_celsius(32)} градусам по Цельсию.")
32 градуса по Фаренгейту равны 0,0 градусам по Цельсию.
```

Для этой задачи вам нужно будет изменить приведенные ниже блоки кода и добавить код для получения желаемых результатов.

### Советы по выполнению упражнения

- Используйте **в точности** те же имена переменных, что и в инструкциях, потому что ваши ответы будут автоматически оценены, и тесты, которые оценивают ваши ответы, зависят от того же форматирования или именования переменных, что и в инструкциях.

- **Пожалуйста**,

    - **Не изменяйте имена файлов**. Сделайте все свое редактирование в файле `Exercise-4-зкщидуь-1.ipynb` (этом файле).
    - **Не копируйте/вставляйте ячейки в этом блокноте**. Мы используем автоматизированную систему оценки, которая не будет работать, если есть копии кодовых ячеек.
    - **Не изменяйте типы ячеек**. Вы можете добавить ячейки, но изменение типов ячеек для существующих ячеек (например, c кода на Markdown) также приведет к выходу из строя автоматизированной проверки.

### Условия успешного выполнения задания

- Создание функции `fahr_to_celsius`
- Определение функции для одного входного параметра `temp_fahrenheit`
- Создание переменной `converted_temp` внутри функции, которой вы должны назначить результат преобразования (то есть входная температура по Фаренгейту, преобразованная в градусы по Цельсию)
- Возвращение преобразованного значения из функции обратно пользователю
- Ответы на вопросы о функциях в конце этого задания
- Добавление комментариев в ваш код и написание блока Docstring, который объясняет, как использовать вашу функцию `fahr_to_celsius` (то есть, описание назначения функции, параметров и возвращаемых значений)
- Загрузка вашего блокнота в репозиторий

## Соглашение об использовании ИИ

**Введите свое имя в ячейку ниже** чтобы подтвердить, что вы следовали [принципам курса об использовании инструментов ИИ](https://geo-python-2024.readthedocs.io/ru/latest/course-info/ai-tools.html) и понимаете, что использование инструментов ИИ считается обманом.

<!-- grade_id: E4-P1-AI-LLM_statement, points: 0, grade: true, solution: true -->
Ваш ответ

## Часть 1

Создайте функцию в ячейке ниже, чтобы преобразовать температуры из Фаренгейта в Цельсий.

- Функция должна называться `fahr_to_celsius`
- У неё должен быть один параметр, вызванный `temp_fahrenheit`
- Она должна возвращать одно значение: преобразованную температуру
- Вы должны добавить комментарии и Docstring

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius, points: 0, grade: false, solution: true -->
```python
# Ваш код здесь
raise NotImplementedError()
```

Используйте свою недавно созданную функцию, чтобы вывести ответы на следующие вопросы:

```python
# 1. Сколько будет 48° по Фаренгейту в Цельсиях?

# 2. Как насчет 71° Фаренгейту в Цельсиях?

```

### Проверки для задачи 1

Запустите кодовые ячейки ниже, чтобы проверить свой ответ. Обратите внимание, что не допускается изменять тестовые ячейки. Ваше решение формально верное, если вы не видите ошибок при запуске этих ячеек. Если вы видите ошибку, вам нужно что-то поправить в ячейках с кодом выше.

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test1, points: 1, grade: true, solution: false -->
```python
import inspect

# Проверка что функция существует
assert inspect.isfunction(fahr_to_celsius), 'Fahr_to_celsius should be a function.'
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test2, points: 0.5, grade: true, solution: false -->
```python
# Убедитесь, что функция имеет один параметр, а имя параметра верно
params = list(inspect.signature(fahr_to_celsius).parameters.keys())
assert len(params) == 1, 'The function should have one parameter'
assert params[0] == 'temp_fahrenheit', 'The parameter name should be "temp_fahrenheit".'
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test3, points: 0.5, grade: true, solution: false -->
```python
# Убедитесь, что функция дает правильные ответы для:
# 1. Сколько будет 48° по Фаренгейту в Цельсиях?
assert round(fahr_to_celsius(48), 2) == 8.89
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test4, points: 0.5, grade: true, solution: false -->
```python
# 2. Сколько будет 71° Фаренгейту в Цельсиях?
assert round(fahr_to_celsius(71), 2) == 21.67
```

### Проверьте свой код

- Убедитесь, что вы использовали заданные имена переменных
- Убедитесь, что вы добавили необходимые комментарии в свой код
- Убедитесь, что у вашей функции есть Docstring, который описывает, что она делает

## Часть 2 

Ответьте на следующие вопросы.

1. Вам ясна концепция функции? Если нет, то что не до конца понятно?
2. Каковы преимущества использования функций?
  
Напишите свои ответы ниже (**активируйте ячейку двойным щелчком**).

<!-- grade_id: problem_1_mg_p2_answers, points: 0.5, grade: true, solution: true -->
Ваш ответ здесь

#### Готово!

Теперь вы готовы продолжить с [задачей 2](Exercise-4-problem-2.ipynb).