## Задача 1 - Пакетная обработка файлов данных с помощью цикла `for`

Эта задача предназначена для моделирования общей проблемы, связанной с файлами данных: их пакетная обработка. Обработка обычно включает процесс работы с несколькими или многими файлами данных, и задачей является генерирование списка имен файлов, которые будут обработаны или сохранены на диск.

Для этой проблемы вам нужно будет изменить уже существующие блоки кода, а также добавить код для получения желаемых результатов.



### Советы по выполнению упражнения

- Используйте **в точности** те же имена переменных, что и в инструкциях, потому что ваши ответы будут автоматически оценены, и тесты, которые оценивают ваши ответы, зависят от того же форматирования или именования переменных, что и в инструкциях.

- **Пожалуйста**,

    - **Не изменяйте имена файлов**. Сделайте все свое редактирование в файле `Exercise-2.ipynb` (этом файле).
    - **Не копируйте/вставляйте ячейки в этом блокноте**. Мы используем автоматизированную систему оценки, которая не будет работать, если есть копии кодовых ячеек.
    - **Не изменяйте типы ячеек**. Вы можете добавить ячейки, но изменение типов ячеек для существующих ячеек (например, c кода на Markdown) также приведет к выходу из строя автоматизированной проверки.

### Условия успешного выполнения задания

 - Создание и использование переменных для получения желаемого текстового формата
 - Использование цикла `for` для итерации по набору чисел
 - Создание требуемых имен файлов
 - Наличие комментариев в коде
 - Ответы на вопросы в конце задания
 - Загрузка вашего блокнота в репозиторий

## Соглашение об использовании ИИ

**Введите свое имя в ячейку ниже** чтобы подтвердить, что вы следовали [принципам курса об использовании инструментов ИИ](https://geo-python-2024.readthedocs.io/ru/latest/course-info/ai-tools.html) и понимаете, что использование инструментов ИИ считается обманом.

<!-- grade_id: E3-P1-AI-LLM_statement, points: 0, grade: true, solution: true -->
Ваш ответ здесь

### Часть 1

- Создайте переменную `basename` которая содержит текст `"Station"`.

<!-- grade_id: problem_1_ag_p1_station_name, points: 0, grade: false, solution: true -->
```python
# Ваш код здесь
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p1_station_name_test1, points: 0.2, grade: true, solution: false -->
```python
# Тест печати переменной
print(basename)
```


- Создайте новую переменную `filenames` как пустой список.

<!-- grade_id: problem_1_ag_p1_empty_filenames_list, points: 0, grade: false, solution: true -->
```python
# Ваш код здесь
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p1_empty_filenames_list_test1, points: 0.3, grade: true, solution: false -->
```python
# Тест печати переменной
print(filenames)
```

### Часть 2

- С помощью цикла `for`, пробегите по числовому диапазону 0-20. В теле цикла вы должны:
    - Создать переменную `station` которая содержит:
        1. текст из переменной `basename`
        2. число
        3. расширение файла `.txt`  
    - Добавьте содержимое `station` в список `filenames`, чтобы получить следующий результат:

      ```
      ['Station_0.txt', 'Station_1.txt', 'Station_2.txt', 'Station_3.txt',
       'Station_4.txt', 'Station_5.txt', 'Station_6.txt', 'Station_7.txt',
       'Station_8.txt', 'Station_9.txt', 'Station_10.txt', 'Station_11.txt',
       'Station_12.txt', 'Station_13.txt', 'Station_14.txt', 'Station_15.txt',
       'Station_16.txt', 'Station_17.txt', 'Station_18.txt', 'Station_19.txt',
       'Station_20.txt']
      ```

<!-- grade_id: problem_1_ag_p2_populate_filenames_list, points: 0, grade: false, solution: true -->
```python
# Ваш код здесь
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p2_populate_filenames_list_test1, points: 0.2, grade: true, solution: false -->
```python
# Тест печати результатов
print(filenames)
```

<!-- grade_id: problem_1_ag_p2_populate_filenames_list_test2, points: 0.3, grade: true, solution: false -->
```python
# Убеждаемся, что значение последней станции верно
assert station.lower().strip() == 'station_20.txt', 'Значение последней станции неверно'
```

<!-- grade_id: problem_1_ag_p2_populate_filenames_list_test3, points: 1, grade: true, solution: false -->
```python
# Убеждаемся, что в списке 21 значение
assert len(filenames) == 21, 'Длина списка "filenames" должна быть 21'
```

### Часть 3

Ниже несколько вопросов, чтобы убедиться, что вы усвоили материал. Пожалуйста, ответьте кратко, в нескольких предложениях, используя ячейку Markdown ниже.

Вы также можете написать любые отзывы или вопросы, касающиеся этого упражнения.

1. Ясна ли вам концепция цикла `for`? Если нет, что было трудно понять?
2. Вы добавили комментарии к своему кода? Если нет, пожалуйста, добавьте сейчас :)
  
Напишите свои ответы ниже (дважды щелкните, чтобы активировать и редактировать ячейку ниже).

<!-- grade_id: problem_1_mg_p3_style_assessment, points: 0, grade: true, solution: true -->
Ваш ответ здесь

### Готово!

Продолжайте работу над [заданием 2](Exercise-3-problem-2.ipynb).