## Задача 2 - Температурный классификатор

Во [второй части третьего упражнения на прошлой неделе](https://gitflic.ru/project/geopython-spbu/exercise-3) мы учились классифицировать температуру на четырех разных класса: cold, slippery, warm и hot. Давайте продолжим работать над той же задачей.

Ниже вы должны создать функцию по имени `temp_classifier`, которая принимает значение температуры в Цельсиях, чтобы затем классифицировать её в целочисленные числа 0-3 на основе следующих критериев:

| Возвращаемое значение | Критерии классификации                                                 |
| :----------: | :----------------------------------------------------------------------- |
| 0            | Температура ниже -2 градуса по Цельсию                                    |
| 1            | Температура равна или теплее -2, но менее +2 градуса Цельсия   |
| 2            | Температура равна или теплее, чем +2, но менее чем +15 градусов по Цельсию  |
| 3            | Температура равна или теплее, чем +15 градусов по Цельсию                   |

### Советы по выполнению упражнения

- Используйте **в точности** те же имена переменных, что и в инструкциях, потому что ваши ответы будут автоматически оценены, и тесты, которые оценивают ваши ответы, зависят от того же форматирования или именования переменных, что и в инструкциях.

- **Пожалуйста**,

    - **Не изменяйте имена файлов**. Сделайте все свое редактирование в файле `Exercise-4-зкщидуь-1.ipynb` (этом файле).
    - **Не копируйте/вставляйте ячейки в этом блокноте**. Мы используем автоматизированную систему оценки, которая не будет работать, если есть копии кодовых ячеек.
    - **Не изменяйте типы ячеек**. Вы можете добавить ячейки, но изменение типов ячеек для существующих ячеек (например, c кода на Markdown) также приведет к выходу из строя автоматизированной проверки.

### Условия успешного выполнения задания

- Создание новой функции с именем `temp_classifier`
- Определение одного параметра, называемого `temp_celsius` в вашей функции` temp_classifier`
- Реклассификация входных температур на основе критериев в приведенной выше таблице и возврат значения класса в виде целочисленного числа (то есть 0, 1, 2 или 3)
- Добавление комментариев в ваш код и Docstring, в котором объясняется, как использовать вашу функцию `temp_classifier`
- Загрузка вашего блокнота в репозиторий

## Соглашение об использовании ИИ

**Введите свое имя в ячейку ниже** чтобы подтвердить, что вы следовали [принципам курса об использовании инструментов ИИ](https://geo-python-2024.readthedocs.io/ru/latest/course-info/ai-tools.html) и понимаете, что использование инструментов ИИ считается обманом.

<!-- grade_id: E4-P2-AI-LLM_statement, points: 0, grade: true, solution: true -->
Ваш ответ:

## Часть 1

Создайте функцию в ячейке ниже, чтобы классифицировать температуры, как указано.

- Ваша функция должна называться `temp_classifier`
- У неё должен быть один параметр с именем `temp_celsius`
- Она должна классифицировать температуры для данного входного значения температуры

<!-- grade_id: problem_2_ag_p1_temp_classifier, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Используйте свою недавно созданную функцию, чтобы распечатать ответы на следующие вопросы:

```python
# 1. Какое значение класса для 16,5 градусов по Цельсию?

# 2.Какое значение класса для +2 градусов по Цельсию?

```

### Часть 2 – проверки

Запустите ячейки ниже, чтобы проверить свой ответ. Обратите внимание, что вы не можете менять тестовые ячейки. Ваше решение верно, если вы не видите никаких ошибок при запуске этих ячеек.

<!-- grade_id: problem_2_ag_p1_temp_classifier_test1, points: 0.5, grade: true, solution: false -->
```python
import inspect

# Проверка на существование функции
assert inspect.isfunction(temp_classifier)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test2, points: 0.5, grade: true, solution: false -->
```python
# Проверяем, что функция имеет один параметр, а имя памаметра верно
params = list(inspect.signature(temp_classifier).parameters.keys())
assert len(params) == 1
assert params[0] == 'temp_celsius'
```

Убедитесь, что функция дает правильные ответы для выбранных значений:

<!-- grade_id: problem_2_ag_p1_temp_classifier_test3, points: 0.25, grade: true, solution: false -->
```python
# 1. Какое значение класса для 16,5 градусов (Цельсия)?
assert temp_classifier(16.5) == 3, 'Неправильный класс'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test4, points: 0.25, grade: true, solution: false -->
```python
# 2. Какое значение класса для +2 градусов (Цельсия)?
assert temp_classifier(2) == 2, 'Неправильный класс'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test5, points: 0.25, grade: true, solution: false -->
```python
# 3. Какое значение класса для +1 градусов (Цельсия)?
assert temp_classifier(1) == 1, 'Неправильный класс'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test6, points: 0.25, grade: true, solution: false -->
```python
# 4. Каково значение класса для -5 градусов (Цельсия)?
assert temp_classifier(-5) == 0, 'Неправильный класс'
print("ok :)")
```

## Часть 2

Проверьте свой код.

- Все ли имена переменных правильны?
- В вашем коде достаточно комментариев?
- Есть ли у вашей функции информативный Docstring?

Вы можете добавить любую обратную связь, которая у вас есть, об этой задача в ячейке ниже.

<!-- grade_id: problem_2_manual_check, points: 1, grade: true, solution: true -->
Ваш ответ здесь